**IMPORTANT** <br> <ul> <li> Do **NOT** replace or remove this notebook (ipynb file)! Each cell has unique nbgrader's metadata and ID which, if changed outside the nbgrader, cannot pass the tests. Do **NOT** change the name of the file!</li> <li> To receive any credit, don't forget to **SUBMIT** your notebook when you are done! You can have multiple submissions before the deadline; only the last one is saved, including its timestamp.</li> <li>Before submitting, **Validate** your notebook to check if your codes pass all visible tests. </li> <li>Make sure you fill in any cell with the comment `# your code here`. Remove or comment the command `fail()` (in R), or `raise NotImplementedError` (in Python) and place your code there </li> </ul>

```python
NAME = "Tomohiro Wada"
```

---

```python
## run this cell first!!!
from nose.tools import assert_equal, assert_raises
import warnings
import numpy as np
```

## Question 1

Create a function `convTemp`, which converts temperature from one of the
units Fahrenheit, Celsius and Kelvin, to another from this list. Your function should have 3 input parameters:

- **x** - which can be either a number (int or float), or array 
- **fro** - which is a character parameter, having possible values “C” (default), “F”, or “K” (**Note**: `from` is a reserved word in Python, so we use `fro` for our input parameter)
- **to** - which is a character parameter, having possible values “F” (default), “C”, or “K”

Use the following relations between different units, with $F$ in Fahrenheits, $C$ in Celsius and $K$ in Kelvin:

\begin{align} 
F &= {9\over 5} \cdot C + 32  \newline 
K &= C + 273.15
\end{align}

If `from` and `to` happen to be the same units (i.e. the same character variables from the set {"C", "F" "K"}),
then your function should return the same vector as the input vector, but should previously print out the
warning, which warns a user that these two are the same (so the function is not useful in that case). Import the module `warnings` and use the following command for printing the warning message (do not change this command, not even a single character!)<br>

`warnings.warn("Your 'fro' parameter is the same as your 'to' parameter!")` <br> 

Here are some examples with various inputs and what your function should return. Keep in mind the output
vector should be of the same length as the input vector x. Also, if you don’t specify the values of from or to,
the default values should be used.

<!-- grade_id: cell-ffbcf0badabea285, points: 0, grade: false, solution: true -->
```python
def convTemp(x=[], fro="C", to="F"):
    single_value = False

    if fro == to:
        warnings.warn("Your 'fro' parameter is the same as your 'to' parameter!")
        return x
    elif isinstance(x, (int, float)):
        x = [x]
        single_value = True

    results = []
    if fro == "C" and to == "F":
        results = [(9/5)*temp+32 for temp in x]
    elif fro == "C" and to == "K":
        results = [temp + 273.15 for temp in x]
    elif fro == "F" and to == "C":
        results = [(temp-32)*(5/9) for temp in x]
    elif fro == "F" and to == "K":
        results = [(temp-32)*(5/9) + 273.15 for temp in x]
    elif fro == "K" and to == "C":
        results = [temp - 273.15 for temp in x]
    elif fro == "K" and to == "F":
        results = [(9/5)*(temp-273.15)+32 for temp in x]

    if single_value:
        return results[0]
    else:
        return results
    
    
# your code here
#raise NotImplementedError
            
```

<!-- grade_id: cell-71fec7304690b95c, points: 5, grade: true, solution: false -->
```python
assert convTemp(0) == 32
assert convTemp([0,10,20]) == [32, 50, 68]


```

<!-- grade_id: cell-f54ef113632b0545, points: 4, grade: true, solution: false -->
```python
assert convTemp(104, fro="F", to="C") == 40
assert convTemp([68, -4, 59], fro="F", to="C") == [20, -20, 15]

```

<!-- grade_id: cell-76141357521f9cc2, points: 4, grade: true, solution: false -->
```python
import numpy as np

assert abs(convTemp(30, to="K") - 303.15) < 1e-10
assert max(abs(np.array(convTemp([283.15, 322], fro="K")) - np.array([50, 119.93]))) < 1e-10

```

<!-- grade_id: cell-aafa54aff03ac7a2, points: 4, grade: true, solution: false -->
```python
## check whether the answer is a list of length n is input is a list of length n

assert type(convTemp([58, 104, 32], fro="F", to="C")) == list
assert len(convTemp([58, 104, 32], fro="F", to="C")) == len([58, 104, 32])

```

<!-- grade_id: cell-97e6851dfda45397, points: 4, grade: true, solution: false -->
```python
with warnings.catch_warnings(record=True) as w:
    # trigger a warning
    convTemp(58, fro="F")
    # verify the message
    assert str(w[-1].message) == "Your 'fro' parameter is the same as your 'to' parameter!"
    

```

<!-- grade_id: cell-49510928a0d3b69a, points: 4, grade: true, solution: false -->
```python
with warnings.catch_warnings(record=True) as w:
    # Trigger a warning.
    convTemp(58, to="K", fro="K")
    # Verify the message
    assert str(w[-1].message) == "Your 'fro' parameter is the same as your 'to' parameter!"
 
```

## Question 2

Create a piece-wise function `pwfun()`, defined in the following way:

$$\text{pwfun}(x) = \begin{cases}
-2 x - 2 &, \hspace{1.1cm} x<-1 \\
\hspace{.5cm} 0 &, -1\leq x \leq 1 \\
x^2 - 1 &, \hspace{0.3cm} 1 < x
\end{cases}$$

Its graph is given below. <br>

 <img src="piecewise-fcn.jpg" height="600" width="300"> 

Your python function should be able to accept both a single numeric value (float or int), and a list of numeric values. The output should be either a single numeric or a list, depending on what the input is. **Round your output values to 2 decimal places.** Also, whenever the output entry is an integer (in mathematical sense), it should be converted into an `int` value. For example, if `x = [-3.0, 5.0]`, then `pwfun(x)` should return

`[4,24]`

(**NOT** `[4., 24.]`, whose entries `4.` and `24.` are floats)

Here are some input vectors and what you should get as a corresponding output. 

`pwfun(-1)` <br> 
`0`

`pwfun([-2, 0.5, 3])` <br>
`[2, 0, 8]`

`x = [-3.0, -2.5, -1, 0 , 1, 2, 3.2]  ## note the output of the first two floats`  <br> 
`pwfun(x)` <br>
`[4, 3, 0, 0, 0, 3, 9.24]`


<!-- grade_id: cell-0439f2e4d601c7a1, points: 0, grade: false, solution: true -->
```python
def pwfun(x = []):
    singleValue = False

    if isinstance(x, (int, float)):
            x = [x]
            singleValue = True

    result = []
    for i in x:
        if i < -1:
            result.append(((-2 * i) - 2))
            if result[-1] % 1 == 0:
                result[-1] = int(result[-1])
            else:
                result[-1] = round(result[-1], 2)
        elif i >= -1 and i <= 1:
            result.append(int(0))
        elif i > 1:
            result.append(((i**2) - 1))
            if result[-1] % 1 == 0:
                result[-1] = int(result[-1])
            else:
                result[-1] = round(result[-1], 2)
        i += 1

    if singleValue:
        return result[0]
    else:
        return result
    
# your code here
#raise NotImplementedError


x = [-3.0, -2.5, -1, 0 , 1, 2, 3.2] 
pwfun(x)
```

Output:
```
[4, 3, 0, 0, 0, 3, 9.24]
```

<!-- grade_id: cell-cbf227b7c693f503, points: 2, grade: true, solution: false -->
```python
assert pwfun(-3) == 4
assert pwfun(0.56) == 0
assert pwfun(4) == 15

x = [-6.0, -0.5 , 0.2,  7] 
assert pwfun(x) == [10, 0, 0, 48]

```

<!-- grade_id: cell-91919c2df753cad3, points: 7, grade: true, solution: false -->
```python
"""testing for various input arguments (hidden tests)"""

```

Output:
```
'testing for various input arguments (hidden tests)'
```

<!-- grade_id: cell-cd4da1c1d69f6719, points: 8, grade: true, solution: false -->
```python
x = [-7/2, -2.5, -1, 0 , 1, 2, 3.2] 
assert pwfun(x) == [5, 3, 0, 0, 0, 3, 9.24]

```

<!-- grade_id: cell-29bfcf36dd047ec6, points: 2, grade: true, solution: false -->
```python
## testing whether the entries are of appropriate type

import numpy as np

x = [-3.0, 2.3, -7/2, -2.5, -1.4] 
out = pwfun(x)  ## components with indices 1 and 4 (i.e. 2nd and the last component) 
                ## should be real, while the rest should be integers

assert type(out[0]) == int  ## should be of int type
assert type(out[1]) in [float, np.float64]  ## should be either float or numpy.float64 type
assert type(out[2]) == int
assert type(out[3]) == int
assert type(out[4]) in [float, np.float64]

```

<!-- grade_id: cell-bd01ec82bf310684, points: 6, grade: true, solution: false -->
```python
"""testing whether the answer is correct (hidden test)"""

```

Output:
```
'testing whether the answer is correct (hidden test)'
```